// variables3.rs
// Execute `rustlings hint variables3` or use the `hint` watch subcommand for a hint.

// I AM NOT DONE

fn main() {
    let x: i32 = 10;
    // let x: i8 = 10;
    // let x: i16 = 10;
    // let x: i32 = 10;
    // let x: i64 = 10;
    // let x: i128 = 10;

    println!("Number {}", x);
}
